Altered visual cortex excitatory/inhibitory ratio following transient congenital visual deprivation in humans

Author response:

The following is the authors’ response to the previous reviews.

eLife Assessment

This neuroimaging and electrophysiology study in a small cohort of congenital cataract patients with sight recovery aims to characterize the effects of early visual deprivation on excitatory and inhibitory balance in visual cortex. While contrasting sight-recovery with visually intact controls suggested the existence of persistent alterations in Glx/GABA ratio and aperiodic EEG signals, it provided only incomplete evidence supporting claims about the effects of early deprivation itself. The reported data were considered valuable, given the rare study population. However, the small sample sizes, lack of a specific control cohort and multiple methodological limitations will likely restrict usefulness to scientists working in this particular subfield.

We thank the reviewing editors for their consideration and updated assessment of our manuscript after its first revision.

In order to assess the effects of early deprivation, we included an age-matched, normally sighted control group recruited from the same community, measured in the same scanner and laboratory. This study design is analogous to numerous studies in permanently congenitally blind humans, which typically recruited sighted controls, but hardly ever individuals with a different, e.g. late blindness history. In order to improve the specificity of our conclusions, we used a frontal cortex voxel in addition to a visual cortex voxel (MRS). Analogously, we separately analyzed occipital and frontal electrodes (EEG).

Moreover, we relate our findings in congenital cataract reversal individuals to findings in the literature on permanent congenital blindness. Note, there are, to the best of our knowledge, neither MRS nor resting-state EEG studies in individuals with permanent late blindness.

Our participants necessarily have nystagmus and low visual acuity due to their congenital deprivation phase, and the existence of nystagmus is a recruitment criterion to diagnose congenital cataracts.

It might be interesting for future studies to investigate individuals with transient late blindness. However, such a study would be ill-motivated had we not found differences between the most “extreme” of congenital visual deprivation conditions and normally sighted individuals (analogous to why earlier research on permanent blindness investigated permanent congenitally blind humans first, rather than permanently late blind humans, or both in the same study). Any result of these future work would need the reference to our study, and neither results in these additional groups would invalidate our findings.

Since all our congenital cataract reversal individuals by definition had visual impairments, we included an eyes closed condition, both in the MRS and EEG assessment. Any group effect during the eyes closed condition cannot be due to visual acuity deficits changing the bottom-up driven visual activation.

As we detail in response to review 3, our EEG analyses followed the standards in the field.

Public Reviews:

Reviewer #1 (Public review):

Summary

In this human neuroimaging and electrophysiology study, the authors aimed to characterise effects of a period of visual deprivation in the sensitive period on excitatory and inhibitory balance in the visual cortex. They attempted to do so by comparing neurochemistry conditions ('eyes open', 'eyes closed') and resting state, and visually evoked EEG activity between ten congenital cataract patients with recovered sight (CC), and ten age-matched control participants (SC) with normal sight.

First, they used magnetic resonance spectroscopy to measure in vivo neurochemistry from two locations, the primary location of interest in the visual cortex, and a control location in the frontal cortex. Such voxels are used to provide a control for the spatial specificity of any effects, because the single-voxel MRS method provides a single sampling location. Using MR-visible proxies of excitatory and inhibitory neurotransmission, Glx and GABA+ respectively, the authors report no group effects in GABA+ or Glx, no difference in the functional conditions 'eyes closed' and 'eyes open'. They found an effect of group in the ratio of Glx/GABA+ and no similar effect in the control voxel location. They then perform multiple exploratory correlations between MRS measures and visual acuity, and report a weak positive correlation between the 'eyes open' condition and visual acuity in CC participants.

The same participants then took part in an EEG experiment. The authors selected two electrodes placed in the visual cortex for analysis and report a group difference in an EEG index of neural activity, the aperiodic intercept, as well as the aperiodic slope, considered a proxy for cortical inhibition. Control electrodes in the frontal region did not present with the same pattern. They report an exploratory correlation between the aperiodic intercept and Glx in one out of three EEG conditions.

The authors report the difference in E/I ratio, and interpret the lower E/I ratio as representing an adaptation to visual deprivation, which would have initially caused a higher E/I ratio. Although intriguing, the strength of evidence in support of this view is not strong. Amongst the limitations are the low sample size, a critical control cohort that could provide evidence for higher E/I ratio in CC patients without recovered sight for example, and lower data quality in the control voxel. Nevertheless, the study provides a rare and valuable insight into experience-dependent plasticity in the human brain.

Strengths of study

How sensitive period experience shapes the developing brain is an enduring and important question in neuroscience. This question has been particularly difficult to investigate in humans. The authors recruited a small number of sight-recovered participants with bilateral congenital cataracts to investigate the effect of sensitive period deprivation on the balance of excitation and inhibition in the visual brain using measures of brain chemistry and brain electrophysiology. The research is novel, and the paper was interesting and well written.

Limitations

Low sample size. Ten for CC and ten for SC, and further two SC participants were rejected due to lack of frontal control voxel data. The sample size limits the statistical power of the dataset and increases the likelihood of effect inflation.

In the updated manuscript, the authors have provided justification for their sample size by pointing to prior studies and the inherent difficulties in recruiting individuals with bilateral congenital cataracts. Importantly, this highlights the value the study brings to the field while also acknowledging the need to replicate the effects in a larger cohort.

Lack of specific control cohort. The control cohort has normal vision. The control cohort is not specific enough to distinguish between people with sight loss due to different causes and patients with congenital cataracts with co-morbidities. Further data from a more specific populations, such as patients whose cataracts have not been removed, with developmental cataracts, or congenitally blind participants, would greatly improve the interpretability of the main finding. The lack of a more specific control cohort is a major caveat that limits a conclusive interpretation of the results.

In the updated version, the authors have indicated that future studies can pursue comparisons between congenital cataract participants and cohorts with later sight loss.

MRS data quality differences. Data quality in the control voxel appears worse than in the visual cortex voxel. The frontal cortex MRS spectrum shows far broader linewidth than the visual cortex (Supplementary Figures). Compared to the visual voxel, the frontal cortex voxel has less defined Glx and GABA+ peaks; lower GABA+ and Glx concentrations, lower NAA SNR values; lower NAA concentrations. If the data quality is a lot worse in the FC, then small effects may not be detectable.

In the updated version, the authors have added more information that informs the reader of the MRS quality differences between voxel locations. This increases the transparency of their reporting and enhances the assessment of the results.

Because of the direction of the difference in E/I, the authors interpret their findings as representing signatures of sight improvement after surgery without further evidence, either within the study or from the literature. However, the literature suggests that plasticity and visual deprivation drives the E/I index up rather than down. Decreasing GABA+ is thought to facilitate experience dependent remodelling. What evidence is there that cortical inhibition increases in response to a visual cortex that is over-sensitised to due congenital cataracts? Without further experimental or literature support this interpretation remains very speculative.

The updated manuscript contains key reference from non-human work to justify their interpretation.

Heterogeneity in patient group. Congenital cataract (CC) patients experienced a variety of duration of visual impairment and were of different ages. They presented with co-morbidities (absorbed lens, strabismus, nystagmus). Strabismus has been associated with abnormalities in GABAergic inhibition in the visual cortex. The possible interactions with residual vision and confounds of co-morbidities are not experimentally controlled for in the correlations, and not discussed.

The updated document has addressed this caveat.

Multiple exploratory correlations were performed to relate MRS measures to visual acuity (shown in Supplementary Materials), and only specific ones shown in the main document. The authors describe the analysis as exploratory in the 'Methods' section. Furthermore, the correlation between visual acuity and E/I metric is weak, not corrected for multiple comparisons. The results should be presented as preliminary, as no strong conclusions can be made from them. They can provide a hypothesis to test in a future study.

This has now been done throughout the document and increases the transparency of the reporting.

P.16 Given the correlation of the aperiodic intercept with age ("Age negatively correlated with the aperiodic intercept across CC and SC individuals, that is, a flattening of the intercept was observed with age"), age needs to be controlled for in the correlation between neurochemistry and the aperiodic intercept. Glx has also been shown to negatively correlates with age.

This caveat has been addressed in the revised manuscript.

Multiple exploratory correlations were performed to relate MRS to EEG measures (shown in Supplementary Materials), and only specific ones shown in the main document. Given the multiple measures from the MRS, the correlations with the EEG measures were exploratory, as stated in the text, p.16, and in Fig.4. yet the introduction said that there was a prior hypothesis "We further hypothesized that neurotransmitter changes would relate to changes in the slope and intercept of the EEG aperiodic activity in the same subjects." It would be great if the text could be revised for consistency and the analysis described as exploratory.

This has been done throughout the document and increases the transparency of the reporting.

The analysis for the EEG needs to take more advantage of the available data. As far as I understand, only two electrodes were used, yet far more were available as seen in their previous study (Ossandon et al., 2023). The spatial specificity is not established. The authors could use the frontal cortex electrode (FP1, FP2) signals as a control for spatial specificity in the group effects, or even better, all available electrodes and correct for multiple comparisons. Furthermore, they could use the aperiodic intercept vs Glx in SC to evaluate the specificity of the correlation to CC.

This caveat has been addressed. The authors have added frontal electrodes to their analysis, providing an essential regional control for the visual cortex location.

Comments on the latest version:

The authors have made reasonable adjustments to their manuscript that addressed most of my comments by adding further justification for their methodology, essential literature support, pointing out exploratory analyses, limitations and adding key control analyses. Their revised manuscript has overall improved, providing valuable information, though the evidence that supports their claims is still incomplete.

We thank the reviewer for suggesting ways to improve our manuscript and carefully reassessing our revised manuscript.

Reviewer #2 (Public review):

Summary:

The study examined 10 congenitally blind patients who recovered vision through the surgical removal of bilateral dense cataracts, measuring neural activity and neuro chemical profiles from the visual cortex. The declared aim is to test whether restoring visual function after years of complete blindness impacts excitation/inhibition balance in the visual cortex.

Strengths:

The findings are undoubtedly useful for the community, as they contribute towards characterising the many ways in which this special population differs from normally sighted individuals. The combination of MRS and EEG measures is a promising strategy to estimate a fundamental physiological parameter - the balance between excitation and inhibition in the visual cortex, which animal studies show to be heavily dependent upon early visual experience. Thus, the reported results pave the way for further studies, which may use a similar approach to evaluate more patients and control groups.

Weaknesses:

The main methodological limitation is the lack of an appropriate comparison group or condition to delineate the effect of sight recovery (as opposed to the effect of congenital blindness). Few previous studies suggested that Excitation/Inhibition ratio in the visual cortex is increased in congenitally blind patients; the present study reports that E/I ratio decreases instead. The authors claim that this implies a change of E/I ratio following sight recovery. However, supporting this claim would require showing a shift of E/I after vs. before the sight-recovery surgery, or at least it would require comparing patients who did and did not undergo the sight-recovery surgery (as common in the field).

We thank the reviewer for suggesting ways to improve our manuscript and carefully reassessing our revised manuscript.

Since we have not been able to acquire longitudinal data with the experimental design of the present study in congenital cataract reversal individuals, we compared the MRS and EEG results of congenital cataract reversal individuals to published work in congenitally permanent blind individuals. We consider this as a resource saving approach. We think that the results of our cross-sectional study now justify the costs and enormous efforts (and time for the patients who often have to travel long distances) associated with longitudinal studies in this rare population.

There are also more technical limitations related to the correlation analyses, which are partly acknowledged in the manuscript. A bland correlation between GLX/GABA and the visual impairment is reported, but this is specific to the patients group (N=10) and would not hold across groups (the correlation is positive, predicting the lowest GLX/GABA ratio values for the sighted controls - opposite of what is found). There is also a strong correlation between GLX concentrations and the EEG power at the lowest temporal frequencies. Although this relation is intriguing, it only holds for a very specific combination of parameters (of the many tested): only with eyes open, only in the patients group.

Given the exploratory nature of the correlations, we do not base the majority of our conclusions on this analysis. There are no doubts that the reported correlations need replication; however, replication is only possible after a first report. Thus, we hope to motivate corresponding analyses in further studies.

It has to be noted that in the present study significance testing for correlations were corrected for multiple comparisons, and that some findings replicate earlier reports (e.g. effects on EEG aperiodic slope, alpha power, and correlations with chronological age).

Conclusions:

The main claim of the study is that sight recovery impacts the excitation/inhibition balance in the visual cortex, estimated with MRS or through indirect EEG indices. However, due to the weaknesses outlined above, the study cannot distinguish the effects of sight recovery from those of visual deprivation. Moreover, many aspects of the results are interesting but their validation and interpretation require additional experimental work.

We interpret the group differences between individuals tested years after congenital visual deprivation and normally sighted individuals as supportive of the E/I ratio being impacted by congenital visual deprivation. In the absence of a sensitive period for the development of an E/I ratio, individuals with a transient phase of congenital blindness might have developed a visual system indistinguishable from normally sighted individuals. As we demonstrate, this is not so. Comparing the results of congenitally blind humans with those of congenitally permanently blind humans (from previous studies) allowed us to identify changes of E/I ratio, which add to those found for congenital blindness.

We thank the reviewer for the helpful comments and suggestions related to the first submission and first revision of our manuscript. We are keen to translate some of them into future studies.

Reviewer #3 (Public review):

This manuscript examines the impact of congenital visual deprivation on the excitatory/inhibitory (E/I) ratio in the visual cortex using Magnetic Resonance Spectroscopy (MRS) and electroencephalography (EEG) in individuals whose sight was restored. Ten individuals with reversed congenital cataracts were compared to age-matched, normally sighted controls, assessing the cortical E/I balance and its interrelationship and to visual acuity. The study reveals that the Glx/GABA ratio in the visual cortex and the intercept and aperiodic signal are significantly altered in those with a history of early visual deprivation, suggesting persistent neurophysiological changes despite visual restoration.

First of all, I would like to disclose that I am not an expert in congenital visual deprivation, nor in MRS. My expertise is in EEG (particularly in the decomposition of periodic and aperiodic activity) and statistical methods.

Although the authors addressed some of the concerns of the previous version, major concerns and flaws remain in terms of methodological and statistical approaches along with the (over)interpretation of the results. Specific concerns include:

(1 3.1) Response to Variability in Visual Deprivation

Rather than listing the advantages and disadvantages of visual deprivation, I recommend providing at least a descriptive analysis of how the duration of visual deprivation influenced the measures of interest. This would enhance the depth and relevance of the discussion.

Although Review 2 and Review 3 (see below) pointed out problems in interpreting multiple correlational analyses in small samples, we addressed this request by reporting such correlations between visual deprivation history and measured EEG/MRS outcomes.

Calculating the correlation between duration of visual deprivation and behavioral or brain measures is, in fact, a common suggestion. The existence of sensitive periods, which are typically assumed to not follow a linear gradual decline of neuroplasticity, does not necessary allow predicting a correlation with duration of blindness. Daphne Maurer has additionally worked on the concept of “sleeper effects” (Maurer et al., 2007), that is, effects on the brain and behavior by early deprivation which are observed only later in life when the function/neural circuits matures.

In accordance with this reasoning, we did not observe a significant correlation between duration of visual deprivation and any of our dependent variables.

(2 3.2) Small Sample Size

The issue of small sample size remains problematic. The justification that previous studies employed similar sample sizes does not adequately address the limitation in the current study. I strongly suggest that the correlation analyses should not feature prominently in the main manuscript or the abstract, especially if the discussion does not substantially rely on these correlations. Please also revisit the recommendations made in the section on statistical concerns.

In the revised manuscript, we explicitly mention that our sample size is not atypical for the special group investigated, but that a replication of our results in larger samples would foster their impact. We only explicitly mention correlations that survived stringent testing for multiple comparisons in the main manuscript.

Given the exploratory nature of the correlations, we have not based the majority of our claims on this analysis.

(3 3.3) Statistical Concerns

While I appreciate the effort of conducting an independent statistical check, it merely validates whether the reported statistical parameters, degrees of freedom (df), and p-values are consistent. However, this does not address the appropriateness of the chosen statistical methods.

We did not intend for the statcheck report to justify the methods used for statistics, which we have done in a separate section with normality and homogeneity testing (Supplementary Material S9), and references to it in the descriptions of the statistical analyses (Methods, Page 13, Lines 326-329 and Page 15, Lines 400-402).

Several points require clarification or improvement:

(4) Correlation Methods: The manuscript does not specify whether the reported correlation analyses are based on Pearson or Spearman correlation.

The depicted correlations are Pearson correlations. We will add this information to the Methods.

(5) Confidence Intervals: Include confidence intervals for correlations to represent the uncertainty associated with these estimates.

We have added the confidence intervals for all measured correlations to the second revision of our manuscript.

(6) Permutation Statistics: Given the small sample size, I recommend using permutation statistics, as these are exact tests and more appropriate for small datasets.

Our study focuses on a rare population, with a sample size limited by the availability of participants. Our findings provide exploratory insights rather than make strong inferential claims. To this end, we have ensured that our analysis adheres to key statistical assumptions (Shapiro-Wilk as well as Levene’s tests, Supplementary Material S9), and reported our findings with effect sizes, appropriate caution and context.

(7) Adjusted P-Values: Ensure that reported Bonferroni corrected p-values (e.g., p > 0.999) are clearly labeled as adjusted p-values where applicable.

In the revised manuscript, we have changed Figure 4 to say ‘adjusted p,’ which we indeed reported.

(8) Figure 2C

Figure 2C still lacks crucial information that the correlation between Glx/GABA ratio and visual acuity was computed solely in the control group (as described in the rebuttal letter). Why was this analysis restricted to the control group? Please provide a rationale.

Figure 2C depicts the correlation between Glx/GABA+ ratio and visual acuity in the congenital cataract reversal group, not the control group. This is mentioned in the Figure 2 legend, as well as in the main text where the figure is referred to (Page 18, Line 475).

The correlation analyses between visual acuity and MRS/EEG measures were only performed in the congenital cataract reversal group since the sighed control group comprised of individuals with vision in the normal range; thus this analyses would not make sense. Table 1 with the individual visual acuities for all participants, including the normally sighted controls, shows the low variance in the latter group.

For variables in which no apiori group differences in variance were predicted, we performed the correlation analyses across groups (see Supplementary Material S12, S15).

We have now highlighted these motivations more clearly in the Methods of the revised manuscript (Page 16, Lines 405-410).

(9 3.4) Interpretation of Aperiodic Signal

Relying on previous studies to interpret the aperiodic slope as a proxy for excitation/inhibition (E/I) does not make the interpretation more robust.

How to interpret aperiodic EEG activity has been subject of extensive investigation. We cite studies which provide evidence from multiple species (monkeys, humans) and measurements (EEG, MEG, ECoG), including studies which pharmacologically manipulated E/I balance.

Whether our findings are robust, in fact, requires a replication study. Importantly, we analyzed the intercept of the aperiodic activity fit as well, and discuss results related to the intercept.

Quote:

“(3.4) Interpretation of aperiodic signal:

- Several recent papers demonstrated that the aperiodic signal measured in EEG or ECoG is related to various important aspects such as age, skull thickness, electrode impedance, as well as cognition. Thus, currently, very little is known about the underlying effects which influence the aperiodic intercept and slope. The entire interpretation of the aperiodic slope as a proxy for E/I is based on a computational model and simulation (as described in the Gao et al. paper).

Apart from the modeling work from Gao et al., multiple papers which have also been cited which used ECoG, EEG and MEG and showed concomitant changes in aperiodic activity with pharmacological manipulation of the E/I ratio (Colombo et al., 2019; Molina et al., 2020; Muthukumaraswamy & Liley, 2018). Further, several prior studies have interpreted changes in the aperiodic slope as reflective of changes in the E/I ratio, including studies of developmental groups (Favaro et al., 2023; Hill et al., 2022; McSweeney et al., 2023; Schaworonkow & Voytek, 2021) as well as patient groups (Molina et al., 2020; Ostlund et al., 2021).

- The authors further wrote: We used the slope of the aperiodic (1/f) component of the EEG spectrum as an estimate of E/I ratio (Gao et al., 2017; Medel et al., 2020; Muthukumaraswamy & Liley, 2018). This is a highly speculative interpretation with very little empirical evidence. These papers were conducted with ECoG data (mostly in animals) and mostly under anesthesia. Thus, these studies only allow an indirect interpretation by what the 1/f slope in EEG measurements is actually influenced.

Note that Muthukumaraswamy et al. (2018) used different types of pharmacological manipulations and analyzed periodic and aperiodic MEG activity in humans, in addition to monkey ECoG (Muthukumaraswamy & Liley, 2018). Further, Medel et al. (now published as Medel et al., 2023) compared EEG activity in addition to ECoG data after propofol administration. The interpretation of our results are in line with a number of recent studies in developing (Hill et al., 2022; Schaworonkow & Voytek, 2021) and special populations using EEG. As mentioned above, several prior studies have used the slope of the 1/f component/aperiodic activity as an indirect measure of the E/I ratio (Favaro et al., 2023; Hill et al., 2022; McSweeney et al., 2023; Molina et al., 2020; Ostlund et al., 2021; Schaworonkow & Voytek, 2021), including studies using scalp-recorded EEG from humans.

In the introduction of the revised manuscript, we have made more explicit that this metric is indirect (Page 3, Line 91), (additionally see Discussion, Page 24, Lines 644-645, Page 25, Lines 650-657).

While a full understanding of aperiodic activity needs to be provided, some convergent ideas have emerged. We think that our results contribute to this enterprise, since our study is, to the best of our knowledge, the first which assessed MRS measured neurotransmitter levels and EEG aperiodic activity. “

(10) Additionally, the authors state:

"We cannot think of how any of the exploratory correlations between neurophysiological measures and MRS measures could be accounted for by a difference e.g. in skull thickness."

(11) This could be addressed directly by including skull thickness as a covariate or visualizing it in scatterplots, for instance, by representing skull thickness as the size of the dots.

We are not aware of any study that would justify such an analysis.

Our analyses were based on previous findings in the literature.

Since to the best of our knowledge, no evidence exists that congenital cataracts go together with changes in skull thickness, and that skull thickness might selectively modulate visual cortex Glx/GABA+ but not NAA measures, we decided against following this suggestion.

Notably, the neurotransmitter concentration reported here is after tissue segmentation of the voxel region. The tissue fraction was shown to not differ between groups in the MRS voxels (Supplementary Material S4). The EEG electrode impedance was lowered to <10 kOhm in every participant (Methods, Page 13, Line 344), and preparation was identical across groups.

(12 3.5) Problems with EEG Preprocessing and Analysis

Downsampling: The decision to downsample the data to 60 Hz "to match the stimulation rate" is problematic. This choice conflates subsequent spectral analyses due to aliasing issues, as explained by the Nyquist theorem. While the authors cite prior studies (Schwenk et al., 2020; VanRullen & MacDonald, 2012) to justify this decision, these studies focused on alpha (8-12 Hz), where aliasing is less of a concern compared of analyzing aperiodic signal. Furthermore, in contrast, the current study analyzes the frequency range from 1-20 Hz, which is too narrow for interpreting the aperiodic signal as E/I. Typically, this analysis should include higher frequencies, spanning at least 1-30 Hz or even 1-45 Hz (not 20-40 Hz).

As previously mentied in the Methods (Page 15 Line 376) and the previous response, the pop_resample function used by EEGLAB applies an anti-aliasing filter, at half the resampling frequency (as per the Nyquist theorem

https://eeglab.org/tutorials/05_Preprocess/resampling.html). The upper cut off of the low pass filter set by EEGlab prior to down sampling (30 Hz) is still far above the frequency of interest in the current study (1-20 Hz), thus allowing us to derive valid results.

Quote:

“- The authors downsampled the data to 60Hz to "to match the stimulation rate". What is the intention of this? Because the subsequent spectral analyses are conflated by this choice (see Nyquist theorem).

This data were collected as part of a study designed to evoke alpha activity with visual white-noise, which ranged in luminance with equal power at all frequencies from 1-60 Hz, restricted by the refresh rate of the monitor on which stimuli were presented (Pant et al., 2023). This paradigm and method was developed by VanRullen and colleagues (Schwenk et al., 2020; Vanrullen & MacDonald, 2012), wherein the analysis requires the same sampling rate between the presented frequencies and the EEG data. The downsampling function used here automatically applies an anti-aliasing filter (EEGLAB 2019) .”

Moreover, the resting-state data were not resampled to 60 Hz. We have made this clearer in the Methods of the second revision (Page 15, Line 367).

Our consistent results of group differences across all three EEG conditions, thus, exclude any possibility that they were driven by aliasing artifacts.

The expected effects of this anti-aliasing filter can be seen in the attached Author response image 1, showing an example participant’s spectrum in the 1-30 Hz range (as opposed to the 1-20 Hz plotted in the manuscript), clearly showing a 30-40 dB drop at 30 Hz. Any aliasing due to, for example, remaining line noise, would additionally be visible in this figure (as well as Figure 3) as a peak.

Power spectral density of one congenital cataract-reversal (CC) participant in the visual stimulation condition across all channels.

The reduced power at 30 Hz shows the effects of the anti-aliasing filter applied by EEGLAB’s pop_resample function.

As we stated in the manuscript, and in previous reviews, so far there has been no consensus on the exact range of measuring aperiodic activity. We made a principled decision based on the literature (showing a knee in aperiodic fits of this dataset at 20 Hz) (Medel et al., 2023; Ossandón et al., 2023), data quality (possible contamination by line noise at higher frequencies) and the purpose of the visual stimulation experiment (to look at the lower frequency range by stimulating up to 60 Hz, thereby limiting us to quantifying below 30 Hz), that 1-20 Hz would be the fit range in this dataset.

Quote:

“(3) What's the underlying idea of analyzing two separate aperiodic slopes (20-40Hz and 1-19Hz). This is very unusual to compute the slope between 20-40 Hz, where the SNR is rather low.

"Ossandón et al. (2023), however, observed that in addition to the flatter slope of the aperiodic power spectrum in the high frequency range (20-40 Hz), the slope of the low frequency range (1-19 Hz) was steeper in both, congenital cataract-reversal individuals, as well as in permanently congenitally blind humans."

The present manuscript computed the slope between 1-20 Hz. Ossandón et al. as well as Medel et al. (2023) found a “knee” of the 1/f distribution at 20 Hz and describe further the motivations for computing both slope ranges. For example, Ossandón et al. used a data driven approach and compared single vs. dual fits and found that the latter fitted the data better. Additionally, they found the best fit if a knee at 20 Hz was used. We would like to point out that no standard range exists for the fitting of the 1/f component across the literature and, in fact, very different ranges have been used (Gao et al., 2017; Medel et al., 2023; Muthukumaraswamy & Liley, 2018). “

(13) Baseline Removal: Subtracting the mean activity across an epoch as a baseline removal step is inappropriate for resting-state EEG data. This preprocessing step undermines the validity of the analysis. The EEG dataset has fundamental flaws, many of which were pointed out in the previous review round but remain unaddressed. In its current form, the manuscript falls short of standards for robust EEG analysis. If I were reviewing for another journal, I would recommend rejection based on these flaws.

The baseline removal step from each epoch serves to remove the DC component of the recording and detrend the data. This is a standard preprocessing step (included as an option in preprocessing pipelines recommended by the EEGLAB toolbox, FieldTrip toolbox and MNE toolbox), additionally necessary to improve the efficacy of ICA decomposition (Groppe et al., 2009).

In the previous review round, a clarification of the baseline timing was requested, which we added. Beyond this request, there was no mention of the appropriateness of the baseline removal and/or a request to provide reasons for why it might not undermine the validity of the analysis.

Quote:

“- "Subsequently, baseline removal was conducted by subtracting the mean activity across the length of an epoch from every data point." The actual baseline time segment should be specified.

The time segment was the length of the epoch, that is, 1 second for the resting state conditions and 6.25 seconds for the visual stimulation conditions. This has been explicitly stated in the revised manuscript (Page 13, Line 354).”

Prior work in the time (not frequency) domain on event-related potential (ERP) analysis has suggested that the baselining step might cause spurious effects (Delorme, 2023) (although see (Tanner et al., 2016)). We did not perform ERP analysis at any stage. One recent study suggests spurious group differences in the 1/f signal might be driven by an inappropriate dB division baselining method (Gyurkovics et al., 2021), which we did not perform.

Any effect of our baselining procedure on the FFT spectrum would be below the 1 Hz range, which we did not analyze.

Each of the preprocessing steps in the manuscript match pipelines described and published in extensive prior work. We document how multiple aspects of our EEG results replicate prior findings (Supplementary Material S15, S18, S19), reports of other experimenters, groups and locations, validating that our results are robust.

We therefore reject the claim of methodological flaws in our EEG analyses in the strongest possible terms.

Quote:

“(3.5) Problems with EEG preprocessing and analysis:

- It seems that the authors did not identify bad channels nor address the line noise issue (even a problem if a low pass filter of below-the-line noise was applied).

As pointed out in the methods and Figure 1, we only analyzed data from two occipital channels, O1 and O2 neither of which were rejected for any participant. Channel rejection was performed for the larger dataset, published elsewhere (Ossandón et al., 2023; Pant et al., 2023). As control sites we added the frontal channels FP1 and Fp2 (see Supplementary Material S14)

Neither Ossandón et al. (2023) nor Pant et al. (2023) considered frequency ranges above 40 Hz to avoid any possible contamination with line noise. Here, we focused on activity between 0 and 20 Hz, definitely excluding line noise contaminations (Methods, Page 14, Lines 365-367). The low pass filter (FIR, 1-45 Hz) guaranteed that any spill-over effects of line noise would be restricted to frequencies just below the upper cutoff frequency.

Additionally, a prior version of the analysis used spectrum interpolation to remove line noise; the group differences remained stable (Ossandón et al., 2023). We have reported this analysis in the revised manuscript (Page 14, Lines 364-357).

Further, both groups were measured in the same lab, making line noise (~ 50 Hz) as an account for the observed group effects in the 1-20 Hz frequency range highly unlikely. Finally, any of the exploratory MRS-EEG correlations would be hard to explain if the EEG parameters would be contaminated with line noise.

- What was the percentage of segments that needed to be rejected due to the 120μV criteria? This should be reported specifically for EO & EC and controls and patients.

The mean percentage of 1 second segments rejected for each resting state condition and the percentage of 6.25 long segments rejected in each group for the visual stimulation condition have been added to the revised manuscript (Supplementary Material S10), and referred to in the Methods on Page 14, Lines 372-373).

- The authors downsampled the data to 60Hz to "to match the stimulation rate". What is the intention of this? Because the subsequent spectral analyses are conflated by this choice (see Nyquist theorem).

This data were collected as part of a study designed to evoke alpha activity with visual white-noise, which changed in luminance with equal power at all frequencies from 1-60 Hz, restricted by the refresh rate of the monitor on which stimuli were presented (Pant et al., 2023). This paradigm and method was developed by VanRullen and colleagues (Schwenk et al., 2020; VanRullen & MacDonald, 2012), wherein the analysis requires the same sampling rate between the presented frequencies and the EEG data. The downsampling function used here automatically applies an anti-aliasing filter (EEGLAB 2019) .

- "Subsequently, baseline removal was conducted by subtracting the mean activity across the length of an epoch from every data point." The actual baseline time segment should be specified.

The time segment was the length of the epoch, that is, 1 second for the resting state conditions and 6.25 seconds for the visual stimulation conditions. This has now been explicitly stated in the revised manuscript (Page 14, Lines 379-380).

- "We excluded the alpha range (8-14 Hz) for this fit to avoid biasing the results due to documented differences in alpha activity between CC and SC individuals (Bottari et al., 2016; Ossandón et al., 2023; Pant et al., 2023)." This does not really make sense, as the FOOOF algorithm first fits the 1/f slope, for which the alpha activity is not relevant.

We did not use the FOOOF algorithm/toolbox in this manuscript. As stated in the Methods, we used a 1/f fit to the 1-20 Hz spectrum in the log-log space, and subtracted this fit from the original spectrum to obtain the corrected spectrum. Given the pronounced difference in alpha power between groups (Bottari et al., 2016; Ossandón et al., 2023; Pant et al., 2023), we were concerned it might drive differences in the exponent values. Our analysis pipeline had been adapted from previous publications of our group and other labs (Ossandón et al., 2023; Voytek et al., 2015; Waschke et al., 2017).

We have conducted the analysis with and without the exclusion of the alpha range, as well as using the FOOOF toolbox both in the 1-20 Hz and 20-40 Hz ranges (Ossandón et al., 2023). The findings of a steeper slope in the 1-20 Hz range as well as lower alpha power in CC vs SC individuals remained stable. In Ossandón et al., the comparison between the piecewise fits and FOOOF fits led the authors to use the former, as it outperformed the FOOOF algorithm for their data.

- The model fits of the 1/f fitting for EO, EC, and both participant groups should be reported.

In Figure 3 of the manuscript, we depicted the mean spectra and 1/f fits for each group.

In the revised manuscript, we added the fit quality metrics (average R2 values > 0.91 for each group and condition) (Methods Page 15, Lines 395-396; Supplementary Material S11) and additionally show individual subjects’ fits (Supplementary Material S11). “

(14) The authors mention:

"The EEG data sets reported here were part of data published earlier (Ossandón et al., 2023; Pant et al., 2023)." Thus, the statement "The group differences for the EEG assessments corresponded to those of a larger sample of CC individuals (n=38) " is a circular argument and should be avoided."

The authors addressed this comment and adjusted the statement. However, I do not understand, why not the full sample published earlier (Ossandón et al., 2023) was used in the current study?

The recording of EEG resting state data stated in 2013, while MRS testing could only be set up by the second half of 2019. Moreover, not all subjects who qualify for EEG recording qualify for being scanned (e.g. due to MRI safety, claustrophobia)

References

Bottari, D., Troje, N. F., Ley, P., Hense, M., Kekunnaya, R., & Röder, B. (2016). Sight restoration after congenital blindness does not reinstate alpha oscillatory activity in humans. Scientific Reports. https://doi.org/10.1038/srep24683

Colombo, M. A., Napolitani, M., Boly, M., Gosseries, O., Casarotto, S., Rosanova, M., Brichant, J. F., Boveroux, P., Rex, S., Laureys, S., Massimini, M., Chieregato, A., & Sarasso, S. (2019). The spectral exponent of the resting EEG indexes the presence of consciousness during unresponsiveness induced by propofol, xenon, and ketamine. NeuroImage, 189(September 2018), 631–644. https://doi.org/10.1016/j.neuroimage.2019.01.024

Delorme, A. (2023). EEG is better left alone. Scientific Reports, 13(1), 2372. https://doi.org/10.1038/s41598-023-27528-0

Favaro, J., Colombo, M. A., Mikulan, E., Sartori, S., Nosadini, M., Pelizza, M. F., Rosanova, M., Sarasso, S., Massimini, M., & Toldo, I. (2023). The maturation of aperiodic EEG activity across development reveals a progressive differentiation of wakefulness from sleep. NeuroImage, 277. https://doi.org/10.1016/J.NEUROIMAGE.2023.120264

Gao, R., Peterson, E. J., & Voytek, B. (2017). Inferring synaptic excitation/inhibition balance from field potentials. NeuroImage, 158(March), 70–78. https://doi.org/10.1016/j.neuroimage.2017.06.078

Groppe, D. M., Makeig, S., & Kutas, M. (2009). Identifying reliable independent components via split-half comparisons. NeuroImage, 45(4), 1199–1211. https://doi.org/10.1016/j.neuroimage.2008.12.038

Gyurkovics, M., Clements, G. M., Low, K. A., Fabiani, M., & Gratton, G. (2021). The impact of 1/f activity and baseline correction on the results and interpretation of time-frequency analyses of EEG/MEG data: A cautionary tale. NeuroImage, 237. https://doi.org/10.1016/j.neuroimage.2021.118192

Hill, A. T., Clark, G. M., Bigelow, F. J., Lum, J. A. G., & Enticott, P. G. (2022). Periodic and aperiodic neural activity displays age-dependent changes across early-to-middle childhood. Developmental Cognitive Neuroscience, 54, 101076. https://doi.org/10.1016/J.DCN.2022.101076

Maurer, D., Mondloch, C. J., & Lewis, T. L. (2007). Sleeper effects. In Developmental Science. https://doi.org/10.1111/j.1467-7687.2007.00562.x

McSweeney, M., Morales, S., Valadez, E. A., Buzzell, G. A., Yoder, L., Fifer, W. P., Pini, N., Shuffrey, L. C., Elliott, A. J., Isler, J. R., & Fox, N. A. (2023). Age-related trends in aperiodic EEG activity and alpha oscillations during early- to middle-childhood. NeuroImage, 269, 119925. https://doi.org/10.1016/j.neuroimage.2023.119925

Medel, V., Irani, M., Crossley, N., Ossandón, T., & Boncompte, G. (2023). Complexity and 1/f slope jointly reflect brain states. Scientific Reports, 13(1), 21700. https://doi.org/10.1038/s41598-023-47316-0

Molina, J. L., Voytek, B., Thomas, M. L., Joshi, Y. B., Bhakta, S. G., Talledo, J. A., Swerdlow, N. R., & Light, G. A. (2020). Memantine Effects on Electroencephalographic Measures of Putative Excitatory/Inhibitory Balance in Schizophrenia. Biological Psychiatry: Cognitive Neuroscience and Neuroimaging, 5(6), 562–568. https://doi.org/10.1016/j.bpsc.2020.02.004

Muthukumaraswamy, S. D., & Liley, D. T. (2018). 1/F electrophysiological spectra in resting and drug-induced states can be explained by the dynamics of multiple oscillatory relaxation processes. NeuroImage, 179(November 2017), 582–595. https://doi.org/10.1016/j.neuroimage.2018.06.068

Ossandón, J. P., Stange, L., Gudi-Mindermann, H., Rimmele, J. M., Sourav, S., Bottari, D., Kekunnaya, R., & Röder, B. (2023). The development of oscillatory and aperiodic resting state activity is linked to a sensitive period in humans. NeuroImage, 275, 120171. https://doi.org/10.1016/J.NEUROIMAGE.2023.120171

Ostlund, B. D., Alperin, B. R., Drew, T., & Karalunas, S. L. (2021). Behavioral and cognitive correlates of the aperiodic (1/f-like) exponent of the EEG power spectrum in adolescents with and without ADHD. Developmental Cognitive Neuroscience, 48, 100931. https://doi.org/10.1016/j.dcn.2021.100931

Pant, R., Ossandón, J., Stange, L., Shareef, I., Kekunnaya, R., & Röder, B. (2023). Stimulus-evoked and resting-state alpha oscillations show a linked dependence on patterned visual experience for development. NeuroImage: Clinical, 103375. https://doi.org/10.1016/J.NICL.2023.103375

Schaworonkow, N., & Voytek, B. (2021). Longitudinal changes in aperiodic and periodic activity in electrophysiological recordings in the first seven months of life. Developmental Cognitive Neuroscience, 47. https://doi.org/10.1016/j.dcn.2020.100895

Schwenk, J. C. B., VanRullen, R., & Bremmer, F. (2020). Dynamics of Visual Perceptual Echoes Following Short-Term Visual Deprivation. Cerebral Cortex Communications, 1(1). https://doi.org/10.1093/TEXCOM/TGAA012

Tanner, D., Norton, J. J. S., Morgan-Short, K., & Luck, S. J. (2016). On high-pass filter artifacts (they’re real) and baseline correction (it’s a good idea) in ERP/ERMF analysis. Journal of Neuroscience Methods, 266, 166–170. https://doi.org/10.1016/j.jneumeth.2016.01.002

Vanrullen, R., & MacDonald, J. S. P. (2012). Perceptual echoes at 10 Hz in the human brain. Current Biology. https://doi.org/10.1016/j.cub.2012.03.050

Voytek, B., Kramer, M. A., Case, J., Lepage, K. Q., Tempesta, Z. R., Knight, R. T., & Gazzaley, A. (2015). Age-related changes in 1/f neural electrophysiological noise. Journal of Neuroscience, 35(38). https://doi.org/10.1523/JNEUROSCI.2332-14.2015

Waschke, L., Wöstmann, M., & Obleser, J. (2017). States and traits of neural irregularity in the age-varying human brain. Scientific Reports 2017 7:1, 7(1), 1–12. https://doi.org/10.1038/s41598-017-17766-4

Changes from the earlier version to the revised version

Abstract

Abstract

Non-human animal models have indicated that the ratio of excitation to inhibition (E/I) in neural circuits is experience dependent and changes across development. Here, we assessed 3T Magnetic Resonance Spectroscopy (MRS) and electroencephalography (EEG) markers of cortical E/I ratio in ten individuals who had been treated for dense bilateral congenital cataracts, after an average of 12 years of blindness, to test for dependence of the E/I ratio in humans on early visual experience. First, participants underwent MRS scanning at rest with their eyes opened and eyes closed, to obtain visual cortex Gamma-Aminobutyric Acid (GABA+) concentration, Glutamate/Glutamine (Glx) concentration and the concentration ratio of Glx/GABA+, as measures of inhibition, excitation, and E/I ratio respectively. Subsequently, EEG was recorded to assess aperiodic activity (1-20 Hz) as a neurophysiological measure of the cortical E/I ratio, during rest with eyes open and eyes closed, and during flickering stimulation. Across conditions, congenital cataract-reversal individuals demonstrated a significantly lower visual cortex Glx/GABA+ ratio, and a higher intercept and steeper aperiodic slope at occipital electrodes, compared to age-matched sighted controls. In the congenital cataract-reversal group, a lower Glx/GABA+ ratio was associated with better visual acuity, and Glx concentration correlated positively with the aperiodic intercept in the conditions with visual input. We speculate that these findings result from an increased E/I ratio of the visual cortex as a consequence of congenital blindness, which might require commensurately increased inhibition in order to balance the additional excitation from restored visual input. The lower E/I ratio in congenital cataract-reversal individuals would thus be a consequence of homeostatic plasticity.

Introduction

Sensitive periods are epochs during the lifespan within which effects of experience on the brain are particularly strong (Knudsen, 2004). Non-human animal work has established that structural remodelling (Bourgeois, 1997) and the development of local inhibitory neural circuits strongly link to the timing of sensitive periods (Gianfranceschi et al., 2003; Hensch et al., 1998; Hensch & Bilimoria, 2012; Hensch & Fagiolini, 2004; Takesian & Hensch, 2013). Early visual experience has been shown to fine-tune local inhibitory circuits (Benevento et al., 1992; Chattopadhyaya et al., 2004; Gandhi et al., 2008; Toyoizumi et al., 2013), which dynamically control feedforward excitation (Tao & Poo, 2005; Wu et al., 2022). The end of the sensitive period has been proposed to coincide with the maturation of inhibitory neural circuits (Hensch, 2005; Wong-Riley, 2021; H. Zhang et al., 2018). Within this framework, neural circuit stability following sensitive periods is maintained via a balance between excitatory and inhibitory transmission across multiple spatiotemporal scales (Froemke, 2015; Haider et al., 2006; Maffei et al., 2004; Takesian & Hensch, 2013; Wu et al., 2022). Such an excitatory/inhibitory (E/I) ratio has been studied at different organizational levels, including the synaptic and neuronal levels, as well as for neural circuits (Van Vreeswijk & Sompolinsky, 1996; Wu et al., 2022).

The experience-dependence of local inhibitory circuit tuning in early development is supported by a large body of work in non-human animals. In particular, studies of the mouse visual cortex has demonstrated a disrupted tuning of local inhibitory circuits as a consequence of lacking visual experience at birth (Hensch & Fagiolini, 2004; Levelt & Hübener, 2012). In addition, dark-reared mice have been shown to have increased spontaneous neural firing in adulthood (Benevento et al., 1992) and a reduced magnitude of inhibition, particularly in layers II/III of the visual cortex (Morales et al., 2002), suggesting an overall higher level of excitation.

Human neuroimaging studies have similarly demonstrated that visual experience during the first weeks and months of life is crucial for the development of visual neural circuits (Baroncelli et al., 2011; Lewis & Maurer, 2005; Maurer & Hensch, 2012; Röder et al., 2021; Röder & Kekunnaya, 2021; Singh et al., 2018). As studies manipulating visual experience are impossible in human research, much of our understanding of the experience-dependence of visual circuit development comes from patients who underwent a transient period of congenital blindness due to dense bilateral congenital cataracts. If human infants born with dense bilateral cataracts are treated later than a few weeks from birth, they suffer from a permanent reduction of visual acuity (Birch et al., 1998; Khanna et al., 2013), stereovision (Birch et al., 1993; Tytla et al., 1993), and impairments in higher-level visual functions such as face perception (Le Grand et al., 2001; Putzar et al., 2010; Röder et al., 2013), coherent motion detection (Bottari et al., 2018; Hadad et al., 2012; Maurer & Lewis, 2017), visual temporal processing (Badde et al., 2020) and visual feature binding (McKyton et al., 2015; Putzar et al., 2007). These visual deficits in congenital cataract-reversal individuals have been attributed to altered neural development due to the absence of early visual experience, as individuals who suffered from developmental cataracts do not typically display a comparable severity of visual impairments (Lewis & Maurer, 2009; Sourav et al., 2020). While the extant literature has reported correlations between structural changes and behavioral outcomes in congenital cataract-reversal individuals (Feng et al., 2021; Guerreiro et al., 2015; Hölig et al., 2023; Pedersini et al., 2023), functional brain imaging (Heitmann et al., 2023; Raczy et al., 2022) and electrophysiological research (Bottari et al., 2016; Ossandón et al., 2023; Pant et al., 2023; Pitchaimuthu et al., 2021) have started to unravel the neural mechanisms which rely on visual experience during early brain development.

Resting-state activity measured via fMRI suggested increased excitation in the visual cortex of congenital cataract-reversal individuals (Raczy et al., 2022): The amplitude of low frequency (<1 Hz) (blood oxygen level-dependent) fluctuations (ALFF) in the visual cortex was increased in congenital cataract-reversal individuals compared to normally sighted controls when they were scanned with their eyes open. Since similar changes were observed in permanently congenitally blind humans, the authors speculated that congenital visual deprivation resulted in an increased E/I ratio of neural circuits due to impaired neural tuning, which was not reinstated after sight restoration (Raczy et al., 2022). Other studies measured resting-state electroencephalogram (EEG) activity and analyzed periodic (alpha oscillations) (Bottari et al., 2016; Ossandón et al., 2023; Pant et al., 2023) as well as aperiodic activity (Ossandón et al., 2023). Both measures pointed towards an higher E/I ratio of visual cortex in congenital cataract-reversal individuals (Ossandón et al., 2023). In recent research, authors have interpreted the slope of the aperiodic component of the EEG power spectral density function as an indirect indication of the relative level of excitation; the flatter the slope, the higher the assumed E/I ratio (R. Gao et al., 2017; Lombardi et al., 2017; McSweeney et al., 2023; Medel et al., 2020; Molina et al., 2020; Muthukumaraswamy & Liley, 2018; Nanda et al., 2023; Schaworonkow & Voytek, 2021). In fact, prospective studies in children have recently reported a flattening of this slope with age, which was interpreted as increasing levels of excitation with age (Favaro et al., 2023; Hill et al., 2022). Ossandón et al. (2023) observed a flatter slope of the aperiodic power spectrum in the high frequency range (20-40 Hz) but a steeper slope of the low frequency range (1-19 Hz) . This pattern was found in both congenital cataract-reversal individuals, as well as in permanently congenitally blind humans. The low frequency range has often been associated with inhibition (Jensen & Mazaheri, 2010; Lozano-Soldevilla, 2018; Lozano-Soldevilla et al., 2014). However, it has remained unclear how to reconcile EEG resting-state findings for lower and higher frequency ranges.

Two studies with permanently congenitally blind humans employed Magnetic Resonance Spectroscopy (MRS) to investigate the concentration of both, the inhibitory neurotransmitter Gamma-Aminobutyric Acid (GABA) and the excitatory neurotransmitters Glutamate/Glutamine (Glx) as proxy measures of visual cortex inhibition and excitation, respectively (Coullon et al., 2015; Weaver et al., 2013). Glutamate/Glutamine concentration was significantly increased in the “visual” cortex of anophthalmic (n = 5) compared to normally sighted individuals, suggesting increased excitability (Coullon et al., 2015). Preliminary evidence in congenitally permanently blind individuals (n = 9) suggested a decreased GABA concentration in the visual cortex compared to normally sighted individuals (Weaver et al., 2013). Thus, these MRS studies corroborated the hypothesis that a lack of visual input at birth enhances relative excitation in visual cortex compared to typical brain development. However, the degree to which neurotransmitter levels recover following sight restoration after a phase of congenital blindness, and how they related to electrophysiological activity, remained unclear.

Here, we filled this gap: we assessed Glutamate/Glutamine (Glx) and Gamma Aminobutyric Acid (GABA+) concentrations using the MEGA-PRESS sequence (Mescher et al., 1998) in individuals whose sight had been restored, on average, after 12 years of congenital blindness. The ratio of Glx/GABA+ concentration was used as a proxy for the ratio of excitatory to inhibitory neurotransmission (Y. Gao et al., 2024; Grent-’t-Jong et al., 2022; Liu et al., 2015; Narayan et al., 2022; Steel et al., 2020; Takei et al., 2016; L. Zhang et al., 2020). Ten congenital cataract-reversal individuals were compared to age-matched, normally sighted controls at rest. In addition to MRS, EEG was recorded to assess and compare aperiodic activity in the same participants. Participants were tested with eyes open and with eyes closed (MRS and EEG), and while viewing visual stimuli (EEG) which changed in luminance (Pant et al., 2023), since both neurotransmitter levels (Kurcyus et al., 2018) and EEG aperiodic activity (Ossandón et al., 2023) systematically varies between these conditions. We predicted an altered visual cortex Glx/GABA+ concentration ratio in the edited MRS signal in congenital cataract-reversal individuals. Since the aperiodic intercept has been linked to broad band neuronal firing (Manning et al., 2009; Musall et al., 2014; Winawer et al., 2013) and based on prior findings suggesting higher excitation in congenital cataract-reversal individuals (Ossandón et al., 2023; Raczy et al., 2022), we predicted a higher intercept as well as an altered slope of the EEG aperiodic component in this group. We further hypothesized that neurotransmitter changes would be concurrent with changes in the slope and intercept of the EEG aperiodic activity in congenital cataract-reversal individuals (Ossandón et al., 2023). Finally, we exploratorily assessed the relationship between the MRS and EEG parameters, as well as their possible link to visual deprivation history and visual acuity in congenital cataract-reversal individuals.

Methods

Participants

We tested two groups of participants. The first group consisted of 10 individuals with a history of dense bilateral congenital cataracts (CC group, 1 female, Mean Age = 25.8 years, Range = 11 – 43.5). Participants in this group were all recruited at the LV Prasad Eye Institute (Hyderabad, India) and the presence of dense bilateral cataracts at birth was confirmed by ophthalmologists and optometrists based on a combination of the following criteria: clinical diagnosis of bilateral congenital cataract, drawing of the pre-surgery cataract, occlusion of the fundus, nystagmus (a typical consequence of congenital visual deprivation), a family history of bilateral congenital cataracts and a visual acuity of fixating and following light (FFL+) or less prior to surgery, barring cases of absorbed lenses. Absorbed lenses occur specifically in individuals with dense congenital cataracts (Ehrlich, 1948), and were diagnosed based on the morphology of the lens, anterior capsule wrinkling, and plaque or thickness of stroma. Prior to cataract surgery, the intactness of the retina is typically checked. Thus, we can exclude major retinal damage as source of group differences.

Duration of deprivation was calculated as the age of the participant when cataract removal surgery was performed on the first eye. Two participants were operated within the first year of life (at 3 months and 9 months of age), all other participants underwent cataract removal surgery after the age of 6 years (Mean Age at Surgery = 11.8 years, SD = 9.7, Range = 0.2 – 31.4). All participants were tested at least 1 year after surgery (Mean Time since Surgery = 14 years, SD = 9.1, Range = 1.8 - 30.9) (Table 1). Visual acuity was significantly below typical vision in this group (Table 1, Supplementary Material S1).

Removed in the revised version: The second group comprised of 10 normally sighted individuals (SC group, 8 males, Mean Age = 26.3 years, Range = 12 – 41.8). Participants across the two groups were age matched (t(9) = −0.12, p = 0.91). Congenital cataract-reversal individuals were clinically screened at the LV Prasad Eye Institute. Both groups did not self-report any neurological or psychiatric conditions, nor any medications.

Earlier version: Additionally, all participants were screened for MRI exclusion criteria using a standard questionnaire from the radiology department. One additional individual was tested in each group; they were excluded from data analysis as their data files were corrupted due to inappropriate file transfer from the scanner. All participants (as well as legal guardians for minors) gave written and informed consent. This study was conducted after approval from the Local Ethical Commission of the LV Prasad Eye Institute (Hyderabad, India) as well as of the Faculty of Psychology and Human Movement, University of Hamburg (Germany).

Revised version: Both groups did not self-report any neurological or psychiatric conditions, nor any medications. Additionally, all participants were screened for MRI exclusion criteria using a standard questionnaire from the radiology department. One additional individual was tested in each group; they were excluded from data analysis as their data files were corrupted due to inappropriate file transfer from the scanner. All participants (as well as legal guardians for minors) gave written and informed consent. This study was conducted after approval from the Local Ethical Commission of the LV Prasad Eye Institute (Hyderabad, India) as well as of the Faculty of Psychology and Human Movement, University of Hamburg (Germany).

Data Collection and Analysis

The present study consisted of three data acquisition parts on the same day: (1) Magnetic Resonance Spectroscopy (MRS, 45-60 min.); (2) Electroencephalography (EEG, 20 min. plus time for capping); (3) visual acuity assessment (3-5 min.).

Magnetic Resonance Spectroscopy

Participants underwent MRI and MRS scanning at LUCID Diagnostics in Hyderabad (India) with a 3T GE SIGNA Pioneer MRI machine employing a 24-channel head coil. An attendant was present in the scanning room for the duration of the scan to ensure that participants were comfortable and followed the instructions.

A T1 weighted whole brain image was collected for each participant (Repetition Time (TR) = 14.97 ms, Echo Time (TE) = 6.74 ms, Matrix size⍰=⍰512⍰×⍰512, In-plane resolution = 0.43⍰×⍰0.43 mm, Slice thickness = 1.6 mm, Axial slices = 188, Interslice interval= -0.8 mm, Inversion time⍰=⍰500 ms, Flip angle⍰=⍰15°). This structural scan enabled registration of every MRS scan to the participants’ anatomical landmarks (Figure 1). For this scan, participants were instructed to stay as still as possible.

Voxel placement for Magnetic Resonance Spectroscopy and electrode placement for Electroencephalography.

a. Position of the frontal cortex (top) and visual cortex (bottom) voxels in a single subject. Skull-stripped figures output from SPM12. b. Electrode montage according to the 10/20 electrode system with marked occipital electrodes preselected for analyses, and frontal electrodes used for control analyses.

Earlier version: The MRS scans consisted of single-voxel spectroscopy data that were collected using the MEGA-PRESS sequence, which allows for in-vivo quantification of the low-concentration metabolites GABA and glutamate+glutamine (Glu+Gln) (Mescher et al., 1998; Mullins et al., 2014). Due to the spectral overlap of GABA (3.0 ppm) and Glu/Gln (3.75 ppm) with the higher concentration peaks of N-Acetyl Aspartate (NAA) and Creatine (Cr), accurate quantification of GABA and Glu/Gln is challenging. MEGA-PRESS uses spectral editing to obtain these measurements. Spectroscopy data consisted of an edit-ON and an edit-OFF spectrum for each voxel, wherein the “ON” and “OFF” refer to whether the frequency of the editing pulse applied is on- or off-resonance with the signal coupled to the GABA complex (applied at approximately 1.9 ppm). Therefore, subtracting repeated acquisitions of the edit-ON and edit-OFF spectra allows for measurement of the magnitude of signals differing in their response to the editing pulse (e.g. GABA), while cancelling out signals that do not (e.g. Cr) (Mescher et al., 1998). Each MEGA-PRESS scan lasted for 8.5 minutes and was acquired with the following specifications: TR = 2000 ms, TE = 68 ms, Voxel size = 40 mm x 30 mm x 25mm, 256 averages. Additionally, 8 unsuppressed water averages were acquired, allowing all metabolites to be referenced to the tissue water concentration. Concentrations of GABA and Glu/Gln quantified from these acquisitions are respectively referred to as GABA+, due to the presence of macromolecular contaminants in the signal (Mullins et al., 2014), and Glx, due to the combined quantification of the Glu, Gln and Glutathione peaks.

Revised version: The MRS scans consisted of single-voxel spectroscopy data that were collected using the MEGA-PRESS sequence, which allows for in-vivo quantification of the low-concentration metabolites GABA and glutamate+glutamine (Glu+Gln) (Mescher et al., 1998; Mullins et al., 2014). Due to the spectral overlap of GABA (3.0 ppm) and Glu/Gln (3.75 ppm) with the higher concentration peaks of N-Acetyl Aspartate (NAA) and Creatine (Cr), accurate quantification of GABA and Glu/Gln is challenging. MEGA-PRESS uses spectral editing to obtain these measurements. Spectroscopy data consisted of an edit-ON and an edit-OFF spectrum for each voxel, wherein the “ON” and “OFF” refer to whether the frequency of the editing pulse applied is on- or off-resonance with the signal coupled to the GABA complex (applied at approximately 1.9 ppm). Therefore, subtracting repeated acquisitions of the edit-ON and edit-OFF spectra allows for measurement of the magnitude of signals differing in their response to the editing pulse (e.g. GABA), while cancelling out signals that do not (e.g. Cr) (Mescher et al., 1998). Each MEGA-PRESS scan lasted for 8.5 minutes and was acquired with the following specifications: TR = 20001ms, TE = 681ms, Voxel size1=140 mm x 30 mm x 25mm, 256 averages. Additionally, 8 unsuppressed water averages were acquired, allowing all metabolites to be referenced to the tissue water concentration. Concentrations of GABA and Glu/Gln quantified from these acquisitions are respectively referred to as GABA+, due to the presence of macromolecular contaminants in the signal (Mullins et al., 2014), and Glx, due to the combined quantification of the Glu, Gln and Glutathione peaks.

Two MRS scans were collected from the visual cortex, centered on the calcarine sulcus of every participant (Figure 1). A prior study with normally sighted individuals suggested that visual cortex Glx and GABA+ concentrations depend on whether the participants were scanned with eyes open or eyes closed (Kurcyus et al., 2018). Therefore, to ensure any group differences were not potentially driven by differences in eye opening/closure, we tested all participants at rest in two conditions – with eyes open (EO) and eyes closed (EC). Both scans were conducted with regular room illumination, that is, without any explicit visual stimulation.

To ensure that we were identifying neurochemical changes specific to visual regions, we selected the frontal cortex as a control region (Figure 1) and collected two scans (EO and EC) from the frontal cortex. The order of the MRS scans was counterbalanced across individuals for both locations and conditions. Two SC subjects did not complete the frontal cortex scan for the EO condition and were excluded from the statistical comparisons of frontal cortex neurotransmitter concentrations. Voxel placement was optimized to avoid the inclusion of the meninges, ventricles, skull and subcortical structures. For each participant a proper placement was ensured by examining the voxel region across the slices in the acquired T1 volume. Saturation bands to nullify the skull signal were placed at the posterior and anterior edge of the visual cortex and frontal cortex voxel, respectively. Due to the limitations of the clinical scanner settings, rotated and skewed voxels were not possible, and therefore voxels were not always located precisely parallel to the calcarine. As documented in Supplementary Material S2, the visual cortex voxel showed significant (>60%) overlap with the V1-V6 region in every individual participant.

MRS Data Analysis

All data analyses were performed in MATLAB (R2018b, The MathWorks Inc.). For MRS data analyses we used Gannet 3.0, a MATLAB based toolbox specialized for the quantification of GABA+ and Glx from edited spectrum data (Edden et al., 2014). Following initial data analysis, all datasets were reanalyzed for quantification of NAA, GABA+ and Glx using linear combination modelling with the Osprey toolbox (v. 2.5.0) (Oeltzschner et al., 2020) in MATLAB 2024a (Supplementary Material S3). Osprey had not been released when the study was originally conceptualized. The results did not differ between analysis toolboxes. Here, we present the originally planned analyses with Gannet 3.0.

GABA+ and edited Glx concentration values were obtained and corrected using the GannetFit, GannetCoRegister,GannetSegment and GannetQuantify functions (Edden et al., 2014). Briefly, the reported water-normalized, alpha-corrected concentration values, were corrected for the differences in GABA concentration and relaxation times between different tissue types in the voxel (grey matter, white matter and cerebrospinal fluid) (Harris et al., 2015). Gannet uses SPM12 to determine the proportion of grey matter, white matter and cerebrospinal fluid in each individual participant’s voxel (Penny et al., 2007). Note that the tissue fraction values did not differ between groups or conditions (all p’s > 0.19, see Supplementary Material S4). GABA+, Glx and Glx/GABA+ values were compared across groups as proxy measures of inhibition, excitation and E/I ratio, respectively. The use of Glx/GABA+ as a proxy measure of E/I neurotransmission is supported by a study that observed a regional balance between Glx and GABA+ at 3T (Steel et al., 2020). Further, the Glx/GABA+ ratio has been employed in prior studies of visual (Takei et al., 2016; L. Zhang et al., 2020), cingulate (Bezalel et al., 2019), frontal (Y. Gao et al., 2024; Liu et al., 2015; Narayan et al., 2022) and auditory cortex (Grent-’t-Jong et al., 2022).

To control for potential unspecified visual cortex changes due to eye pathology, as opposed to genuine changes in neurotransmitter ratio, we compared N-Acetyl Aspartate (NAA) concentrations in the visual cortex of CC vs SC individuals. NAA forms one of the most prominent peaks in the MR spectrum (2.0 ppm chemical shift). NAA has been quantified with high reproducibility in the visual cortex (Brooks et al., 1999) and medial-temporal cortex (Träber et al., 2006) of neuro-typical individuals as well as in various pathologies across visual, frontal and temporal cortex (Paslakis et al., 2014), for example, schizophrenia (Mullins et al., 2003). We did not expect to find differences in NAA concentration between CC and SC individuals as it has not been demonstrated to vary in anophthalmia (Coullon et al., 2015) or permanent early blindness (Weaver et al., 2013) in humans. TARQUIN 4.3.11 was employed to analyze the OFF-spectrum data (Wilson et al., 2011) to assess NAA concentration. FID-A toolbox was used to correct the data for phase errors across acquisitions arising from temporal changes in the magnetic field strength or participant motion (Simpson et al., 2017).

The reported values in the results are water-normalized. All data analyses were repeated with Cr-normalized values from Gannet 3.0, and significant results were replicated (Supplementary Material S5).

MRS Data Quality

The MRS minimum reporting standards form is found in the Supplementary Excel File 1. Mean Signal-to-Noise Ratio values for GABA+ and Glx in all groups and conditions were above 19 in the visual cortex and above 8 in the frontal cortex (Table 2). A recent study has suggested that an SNR value above 3.8 allows for reliable quantification of GABA+ (Zöllner et al., 2021), in conjunction with considering a given study’s sample size (Mikkelsen et al., 2018). Cramer-Rao Lower Bound (CRLB) values, that is, the theoretical lower limit of estimated error, were 30% or lower for NAA quantification in both groups and conditions (Cavassila et al., 2001). Note that CRLB values above 50% are considered unreliable (Wilson et al., 2019). In all quality metrics for Glx, GABA+ and NAA our dataset showed higher quality for the visual cortex voxel than for the frontal cortex voxel, irrespective of group (Main effect of region: all p’s < 0.004, Supplementary Material S6). Such region effects have repeatedly been reported in the MRS literature. They were attributed to magnetic field distortions (Juchem & Graaf, 2017) resulting from the proximity of the frontal cortex voxel to the sinuses. We chose a frontal control voxel rather than a parietal/sensorimotor control voxel (Coullon et al., 2015; Weaver et al., 2013) due to well documented changes in multisensory cortical region as a consequence of congenital blindness (Harrar et al., 2018; Henschke et al., 2017; F. Jiang et al., 2016; Röder et al., 1999; Sabourin et al., 2022; Zatorre et al., 2012) . The fit error for the frontal cortex voxel was below 8.31% for GABA+ and Glx in both groups (Table 2). No absolute cutoffs exist for fit errors. However, Mikkelsen et al. reported a mean GABA+ fit error of 6.24 +/-1.95% from a posterior cingulate cortex voxel across 8 GE scanners using the Gannet pipeline (Mikkelsen et al., 2017). Previous studies in special populations have used frontal cortex data with a fit error of <10% to identify differences between cohorts (Y. Gao et al., 2024; Maier et al., 2022; Pitchaimuthu et al., 2017). Importantly, in the present study, data quality did not significantly differ between groups for GABA+, Glx or NAA (Supplementary Material S6), making it highly unlikely that data quality differences contributed to group differences.

Removed in the revised version: Prior to in-vivo scanning, we confirmed the GABA+ and GABA+/Glx quantification quality with phantom testing (Henry et al., 2011; Jenkins et al., 2019). Imaging sequences were robust in identifying differences of 0.02 mM in GABA concentration; the known vs. measured concentrations of both GABA (r = 0.81, p = 0.004) and GABA/Glx (r = 0.71, p = 0.019) showed significant agreement. This 0.02 mM difference was documented by Weaver et al. (2013) between the occipital cortices of early blind and sighted individuals. The detailed procedure and results are described in Supplementary Material S7. The spectra from all individual subjects are shown in Supplementary Material S8.

MRS Statistical Analysis

All statistical analyses were performed using MATLAB R2018b and R v3.6.3.

We compared the visual cortex concentrations of 3 neurochemicals (GABA+ and Glx from the DIFF spectrum, NAA from the edit-OFF spectrum) between the two groups. For each metabolite we submitted the concentration values from the visual cortices of CC and SC individuals to a group (2 Levels: CC, SC)-by-condition (2 Levels: EO, EC) ANOVA model. To compare the Glx/GABA+ ratio between groups, we additionally submitted this ratio value to a group-by-condition ANOVA. Identical analyses were performed for the corresponding frontal cortex neurotransmitter values. Wherever necessary, post-hoc comparisons were performed using t-tests. The data were tested for normality (Shapiro-Wilk) and homogeneity of variance (Levene’s Test) in R v3.6.3 (Supplementary Material S9). In all ANOVA models, the residuals did not significantly differ from normality.

Electrophysiological recordings

EEG data analyzed in the present study are a subset of datasets that were included in previous reports (Ossandón et al., 2023; Pant et al., 2023). The EEG datasets were re-analyzed to investigate aperiodic activity in the same participants who took part in the MRS study. MRS and EEG data were acquired on the same day. The EEG was recorded in three conditions: (1) at rest with eyes open (EO, 3 minutes), (2) at rest with eyes closed (EC, 3 minutes) and (3) during visual stimulation with stimuli that changed in luminance (LU) with equal power at all frequencies (0-30 Hz) (Pant et al., 2023). We used the slope of the aperiodic (1/f) component of the EEG spectrum as an estimate of E/I ratio (R. Gao et al., 2017; Medel et al., 2020; Muthukumaraswamy & Liley, 2018) and the intercept as an estimate of broadband neuronal firing activity (Haller et al., 2018; Manning et al., 2009; Miller, 2010).

The EEG was recorded using Ag/AgCl electrodes attached according the 10/20 system(Homan et al., 1987) to an elastic cap (EASYCAP GmbH, Herrsching, Germany) (Figure 1). We acquired 32 channel EEG using the BrainAmp amplifier, with a bandwidth of 0.01–200⍰Hz, sampling rate of 5⍰kHz and a time constant of 0.016 Hz /(10 s) (http://www.brainproducts.com/). All scalp recordings were performed against a left ear lobe reference. Electrode impedance was kept below 10 kOhm in all participants.

Participants were asked to sit as still as possible while the EEG was recorded. First, resting-state EEG data were collected. During the EO condition, participants were asked to look towards a blank screen and to avoid eye movements. During the EC condition, participants were instructed to keep their eyes closed. The order of conditions was randomized across participants.

Subsequently, EEG data were recorded during 100 trials of a target detection task with stimuli that changed in luminance (LU). Stimuli were presented with a Dell laptop, on a Dell 22-inch LCD monitor with a refresh rate of 60⍰Hz. They were created with MATLAB r2018b (The MathWorks, Inc., Natick, MA) and the Psychtoolbox 3 (Brainard, 1997; Kleiner et al., 2007). On each trial, participants observed a circle at the center of a black screen, subtending a visual angle of 17 degrees. The circle appeared for 6.25 s and changed in luminance with equal power at all frequencies (0-30 Hz). At the end of every trial, participants had to indicate whether a target square, subtending a visual angle of 6 degrees, appeared on that trial. The experiment was performed in a darkened room (for further details, see (Pant et al., 2023).

EEG Data Analysis

Data analysis was performed using the EEGLab toolbox on MATLAB 2018b (Delorme & Makeig, 2004). All EEG datasets were filtered using a Hamming windowed sinc FIR filter, with a high-pass cutoff at 1 Hz and a low-pass cutoff at 45 Hz. A prior version of the analysis was conducted with line noise removal via spectrum interpolation (Ossandón et al., 2023). However, the analyses reported here did not include this step, since we implemented a low-pass cutoff (20 Hz) which falls far below the typical line noise frequency (50 Hz). Eye movement artifacts were detected in the EEG datasets via independent component analysis using the runica.m function’s Infomax algorithm in EEGLab. Components corresponding to horizontal or vertical eye movements were identified via visual inspection based on criteria discussed in Plöchl et al (Plöchl et al., 2012) and removed.

Earlier version: The two 3 minutes long resting-state recordings (EC, EO) were divided into epochs of 1 s. Epochs with signals exceeding ±120 μV were rejected for all electrodes (see Supplementary Material S10 for percentages by group and condition). We then calculated the power spectral density of the EO and EC resting-state data using the pwelch function (window length = 1000, overlap = 0).

Revised version: The two 3 minutes long resting-state recordings (EC, EO) were divided into epochs of 1 s. Epochs with signals exceeding ±120 μV were rejected for all electrodes (see Supplementary Material S10 for percentages by group and condition). We then calculated the power spectral density of the EO and EC resting-state data using the pwelch function (sampling rate = 1000 Hz, window length = 1000 samples, overlap = 0).

Earlier version: Datasets collected while participants viewed visual stimuli that changed in luminance (LU) were downsampled to 60 Hz (antialiasing filtering performed by EEGLab’s pop_resample function) to match the stimulation rate. The datasets were divided into 6.25 s long epochs corresponding to the duration of visual stimulation per trial. Subsequently, baseline removal was conducted by subtracting the mean activity across the length of an epoch (1 s for the EO and EC conditions, 6.25 s for the LU condition) from every data point. After baseline removal, epochs with signals exceeding a threshold of ±120 μV were rejected in order to exclude potential artifacts. Finally, we calculated the power spectral density of the LU data using the pwelch function (window length = 60 samples, overlap = 0).

Revised version: Datasets collected while participants viewed visual stimuli that changed in luminance (LU) were downsampled to 60⍰Hz (antialiasing filtering performed by EEGLab’s pop_resample function) to match the stimulation rate. The datasets were divided into 6.25 s long epochs corresponding to the duration of visual stimulation per trial. Subsequently, baseline removal was conducted by subtracting the mean activity across the length of an epoch (1 s for the EO and EC conditions, 6.25 s for the LU condition) from every data point. After baseline removal, epochs with signals exceeding a threshold of ±120 μV were rejected in order to exclude potential artifacts. Finally, we calculated the power spectral density of the LU data using the pwelch function (sampling rate = 60 Hz, window length = 60 samples, overlap = 0).

We derived the aperiodic (1/f) component of the power spectrum for the EO, EC and LU conditions (Donoghue, Haller, et al., 2020; Schaworonkow & Voytek, 2021). First, we fit the 1/f distribution function to the frequency spectrum of each participant, separately for each electrode. The 1/f distribution was fit to the normalized spectrum converted to log-log scale (range = 1-20 Hz) (Donoghue, Dominguez, et al., 2020; Gyurkovics et al., 2021; Schaworonkow & Voytek, 2021). We excluded the alpha range (8-14 Hz) for this fit to avoid biasing the results due to documented differences in alpha activity between CC and SC individuals (Bottari et al., 2016; Ossandón et al., 2023; Pant et al., 2023). This 1/f fit resulted in a value of the aperiodic slope, an aperiodic intercept value corresponding to the broadband power of 1-20 Hz, and a fit error value for the spectrum of every participant, individually for each electrode. Spectra from individual subjects are displayed in Supplementary Material S11. The visual cortex aperiodic slope and intercept values were obtained by averaging across the pre-selected occipital electrodes O1 and O2, resulting in one value of broadband slope and one value of intercept per participant and condition (Figure 1). This procedure yielded average R2 values > 0.91 for the aperiodic fit in each group and condition (Supplementary Material S11).

EEG Statistical Analysis

We compared the average visual cortex aperiodic slope and intercept in separate group (2 Levels: CC, SC) by condition (3 levels: EC, EO, LU) ANOVA models. The data were tested for normality (Shapiro-Wilk) and homogeneity of variance (Levene’s Test) in R v3.6.3) (see Supplementary Material S9); in all ANOVA models, the residuals did not significantly differ from normality.

Visual acuity

Earlier version: Visual acuity was measured binocularly for every participant on the date of testing, using the Freiburg Visual Acuity Test (FrACT) (Bach 1996, Bach 2007, https://michaelbach.de/fract/). Visual acuity is reported as the logarithm of the minimum angle of resolution (logMAR, Table 1), wherein higher values indicate worse vision (Elliott, 2016). Analogous to previous studies, we ran a number of exploratory correlation analyses between GABA+, Glx and Glx/GABA+ concentrations and visual acuity at the date of testing, duration of visual deprivation, and time since surgery respectively in the CC group (Birch et al., 2009; Guerreiro et al., 2015; Kalia et al., 2014; Rajendran et al., 2020). We additionally tested the correlation between the neurotransmitter levels and chronological age across the CC and SC groups.

Revised version: Visual acuity was measured binocularly for every participant on the date of testing, using the Freiburg Visual Acuity Test (FrACT) (Bach 1996, Bach 2007, https://michaelbach.de/fract/). Visual acuity is reported as the logarithm of the minimum angle of resolution (logMAR, Table 1), wherein higher values indicate worse vision (Elliott, 2016). Analogous to previous studies, we ran a number of exploratory correlation analyses between GABA+, Glx and Glx/GABA+ concentrations and visual acuity at the date of testing, duration of visual deprivation, and time since surgery, respectively, in the CC group (Birch et al., 2009; Guerreiro et al., 2015; Kalia et al., 2014; Rajendran et al., 2020). As expected from normal vision in the SC group, they did not show considerable variance in visual acuity (Table 1); thus, we refrained from calculation correlations between visual acuity and MRS/EEG parameters in the SC group. Based on the literature, we additionally tested the correlation between the neurotransmitter levels and chronological age across the CC and SC groups. All reported correlation coefficients are Pearson correlations, and 95% confidence intervals were calculated for all correlation coefficients.

Exploratory correlation analyses between MRS and EEG measures

Earlier version: Exploratory correlation analyses between EEG and MRS measures were run separately for CC and SC individuals. We calculated correlations between the aperiodic intercept and GABA+, Glx and Glx/GABA+ concentrations. Further, correlations between the aperiodic slope, and the concentrations of GABA+, Glx and Glx/GABA+ were assessed. MRS measures collected at rest with eyes open (EO) and eyes closed (EC) were correlated with the corresponding resting-state EEG conditions (EO, EC). EEG metrics for the visual stimulation (LU) condition with flickering stimuli were tested for correlation with GABA+, Glx and Glx/GABA+ concentration measured while participants’ eyes were open at rest. We did not have prior hypotheses as to the best of our knowledge no extant literature has tested the correlation between aperiodic EEG activity and MRS measures of GABA+,Glx and Glx/GABA+. Therefore, we corrected for multiple comparisons using the Bonferroni correction (6 comparisons).

Revised version: Exploratory correlation analyses between EEG and MRS measures were run separately for CC and SC individuals. We calculated Pearson correlations between the aperiodic intercept and GABA+, Glx and Glx/GABA+ concentrations. Further, Pearson correlations between the aperiodic slope, and the concentrations of GABA+, Glx and Glx/GABA+ were assessed. MRS measures collected at rest with eyes open (EO) and eyes closed (EC) were correlated with the corresponding resting-state EEG conditions (EO, EC). EEG metrics for the visual stimulation (LU) condition with flickering stimuli were tested for correlation with GABA+, Glx and Glx/GABA+ concentration measured while participants’ eyes were open at rest. We did not have prior hypotheses as to the best of our knowledge no extant literature has tested the correlation between aperiodic EEG activity and MRS measures of GABA+,Glx and Glx/GABA+. Therefore, we corrected for multiple comparisons using the Bonferroni correction (6 comparisons).

Results

Transient congenital visual deprivation lowered the Glx/GABA+ concentration in the visual cortex

The Glx/GABA+ concentration ratio was significantly lower in the visual cortex of congenital cataract-reversal (CC) than age-matched, normally sighted control (SC) individuals (main effect of group: F(1,39) = 5.80, p = 0.021, ηp² = 0.14) (Figure 2). This effect did not vary with eye closure (main effect of condition: F(1,39) = 2.29, p = 0.139, ηp² = 0.06, group-by-condition interaction: F(1,39) = 1.15, p = 0.290, ηp² = 0.03). As a control for unspecific effects of surgery or related to visual deprivation on neurochemistry, the frontal cortex Glx/GABA+ concentration was compared between groups. There was no difference between CC and SC individuals in their frontal cortex Glx/GABA+ concentration (main effect of group: F(1,37) = 0.05, p = 0.82, ηp² < 0.01, main effect of condition: F(1,37) = 2.98, p = 0.093, ηp² = 0.07, group-by-condition interaction: F(1,37) = 0.09, p = 0.76, ηp² < 0.01) (Figure 2).

Edited spectra obtained from Magnetic Resonance Spectroscopy (MRS).

Earlier version: a. Average edited spectra showing GABA+ and edited Glx peaks in the visual cortices of normally sighted individuals (SC, green) and individuals with reversed congenital cataracts (CC, red) are shown. Edited MRS DIFF spectra are separately displayed for the eyes open (EO), and eyes closed (EC) conditions using dashed and solid lines respectively. The standard error of the mean is shaded. Water-normalized GABA+, water-normalized Glx, and Glx/GABA+ concentration distributions for each group and condition are depicted as violin plots on the right. The solid black lines indicate mean values, and dotted lines indicate median values. The coloured lines connect values of individual participants across conditions. b. Corresponding average edited MRS spectra and water-normalized GABA+, water-normalized Glx and Glx/GABA+ concentration distributions measured from the frontal cortex are displayed. c. Correlations between visual cortex Glx/GABA+ concentrations in the visual cortex of CC individuals and visual acuity in logMAR units are depicted for the eyes closed (EC, left) and eyes open (EO, right) conditions.

Revised version: a. Average edited spectra showing GABA+ and edited Glx peaks in the visual cortices of normally sighted individuals (SC, green) and individuals with reversed congenital cataracts (CC, red) are shown. Edited MRS DIFF spectra are separately displayed for the eyes open (EO), and eyes closed (EC) conditions using dashed and solid lines respectively. The standard error of the mean is shaded. Water-normalized GABA+, water-normalized Glx, and Glx/GABA+ concentration distributions for each group and condition are depicted as violin plots on the right. The solid black lines indicate mean values, and dotted lines indicate median values. The coloured lines connect values of individual participants across conditions. b. Corresponding average edited MRS spectra and water-normalized GABA+, water-normalized Glx and Glx/GABA+ concentration distributions measured from the frontal cortex are displayed. c. Correlations between visual cortex Glx/GABA+ concentrations in the visual cortex of CC individuals and visual acuity in logMAR units are depicted for the eyes closed (EC, left) and eyes open (EO, right) conditions. The 95% confidence intervals (CI) of the correlation coefficients (r) are reported.

When separately comparing CC and SC individuals’ GABA+ and Glx concentrations in the visual cortex, we did not find any significant group difference (GABA+ main effect of group: F(1,39) = 2.5, p = 0.12, ηp² = 0.06, main effect of condition: F(1,39) = 0.6, p = 0.43, ηp² = 0.02, group-by-condition interaction: F(1,39) = 0.03, p = 0.86, ηp² < 0.001; Glx main effect of group: F(1,39) = 2.8, p = 0.103, ηp² = 0.07, main effect of condition: F(1,39) = 1.8, p = 0.19, ηp² = 0.05, group-by-condition interaction: F(1,39) = 1.27, p = 0.27, ηp² -0.03) (Figure 2). In the frontal cortex, GABA+ and Glx concentrations did not vary either with group or condition (all p values > 0.19, all ηp² < 0.05) (Figure 2). Note that these findings were replicated when Osprey’s quantification method was used: Glx/GABA+ concentration was lower in the visual cortex of CC than SC individuals, while GABA+ and Glx concentration did not significantly differ (Supplementary Material S3). When analyses were repeated with Cr-normalized Glx/GABA+ concentrations, the Glx concentration was found to be significantly lower in CC vs SC individuals’ visual cortices, in addition to the lower Glx/GABA+ concentration ratio (Supplementary Material S5). Since this finding was not replicated with water normalized Glx concentration in Gannet or Osprey, we refrain from interpreting this additional group effect for Glx.

The Glx/GABA+ concentration measured when CC individuals’ eyes were closed correlated positively with visual acuity on the logMAR scale (r = 0.65, p = 0.044), indicating that CC individuals with higher Glx/GABA+ values had worse visual acuity (Figure 2C, Supplementary Material S12). The same correlation was not significant for the eyes opened condition (r = -0.042, p = 0.908) (Figure 2C). Duration of deprivation and time since surgery did not significantly predict Glx/GABA+, GABA+ or Glx concentrations in the CC group (all p values > 0.088, Supplementary Material S12).

No difference in NAA concentration between CC and SC individuals’ visual cortices.

As a control measure to ensure that between-group differences were specific to hypothesized changes in Glx and GABA+ concentrations, we compared the NAA concentration between CC and SC individuals. The NAA concentration did not significantly differ between groups, neither in visual (main effect of group: F(1,39) = 0.03, p = 0.87, ηp² < 0.001, main effect of condition: F(1,39) = 0.31, p = 0.58, ηp² <0.01, group-by-condition interaction: F(1,39) = 0.09, p = 0.76, ηp² <0.01) nor frontal cortex (main effect of group: F(1,37) = 1.1, p = 0.297, ηp² = 0.02, main effect of condition: F(1,37) = 0.14, p = 0.71, ηp² = 0.01, group-by-condition interaction: F(1,37) = 0.03, p = 0.86, ηp² < 0.001) (Supplementary Material S13, Figure S13).

Transient congenital visual deprivation resulted in a steeper aperiodic slope and higher aperiodic intercept at occipital sites.

The aperiodic slope (1-20 Hz), measured via EEG as an electrophysiological estimate of the E/I ratio (R. Gao et al., 2017; Muthukumaraswamy & Liley, 2018), was compared between CC and SC individuals. The aperiodic slope was significantly steeper i.e. more negative, at occipital electrodes in CC than in SC individuals (F(1,59) = 13.1, p < 0.001, ηp² = 0.19) (Figure 3). Eye closure and visual stimulation did not affect the steepness of the aperiodic slope (F(2,59) = 0.78, p = 0.465, ηp² = 0.03, group-by-condition interaction: F(2,59) = 0.12, p = 0.885, ηp² < 0.01).

Full spectrum and aperiodic activity of the electroencephalogram (EEG).

a. EEG spectra across O1 and O2 with the corresponding aperiodic (1/f) fits for normally sighted individuals (SC, blue, left) and individuals with reversed congenital cataracts (CC, red, right). Spectra of EEG recordings are displayed for the eyes closed (EC) and eyes opened (EO) conditions, as well as while viewing stimuli that changed in luminance (LU). Shaded regions represent the standard error of the mean. b. Aperiodic intercept (top) and slope (bottom) value distributions for each group and condition are displayed as violin plots. Solid black lines indicate mean values, dotted black lines indicate median values. Coloured lines connect values of individual participants across conditions.

The aperiodic intercept (1-20 Hz) was compared between CC and SC individual to estimate group differences in broadband neural activity (Manning et al., 2009; Musall et al., 2014; Winawer et al., 2013) and was found to be significantly larger at occipital electrodes in CC than SC individuals (main effect of group: F(1,59) = 5.2, p = 0.026, ηp² = 0.09) (Figure 3). Eye closure did not affect the magnitude of the aperiodic intercept in either group (main effect of condition: F(2,59) = 0.16, p = 0.848, ηp² < 0.01, group-by-condition interaction: F(2,59) = 0.11, p = 0.892, ηp² < 0.01). No significant group differences in slope and intercept were found for frontal electrodes (Supplementary Material S14).

Within the CC group, visual acuity, time since surgery and duration of blindness did not significantly correlate with the aperiodic slope or the intercept (all p’s > 0.083, Supplementary Material S15). Age negatively correlated with the aperiodic intercept across CC and SC individuals, that is, a reduction of the intercept was observed with age. Similar effects of chronological age have been previously observed (Hill et al., 2022; Voytek et al., 2015) (Supplementary Material S15).

Glx concentration predicted the aperiodic intercept in CC individuals’ visual cortices during ambient and flickering visual stimulation.

We exploratorily tested the relationship between Glx, GABA+ and Glx/GABA+ measured at rest and the EEG aperiodic intercept measured at rest and during flickering visual stimulation, separately for the CC and the SC group. Visual cortex Glx concentration in CC individuals was positively correlated with the aperiodic intercept either when participants had their eyes open during rest (r = 0.91, p = 0.001, Bonferroni corrected) or when they viewed flickering stimuli (r = 0.90, p < 0.001, Bonferroni corrected). Corresponding correlations were not significant for Glx concentrations in the eyes closed condition ( r = 0.341, p > 0.99, Bonferroni corrected). Moreover, in SC individuals no significant correlation was observed between visual cortex Glx concentration and aperiodic intercept in any condition (all p’s > 0.99, Bonferroni corrected) (Figure 4). Given the correlation between the aperiodic intercept and chronological age across groups (Supplementary Material S15), we performed a post-hoc linear regression analysis to model the aperiodic intercept in the CC group with both age and Glx concentration as covariates. Glx concentration, but not age, significantly predicted the aperiodic intercept within the CC group during rest with eyes open and during visual stimulation (Supplementary Material S16).

Exploratory correlation analyses between the aperiodic intercept (1-20 Hz) and glutamate/glutamine (Glx) concentration in the visual cortex.

Earlier version: Correlations between water-normalized Glx concentration and aperiodic intercept are shown for the eyes closed (EC, left), eyes open (EO, middle) and visual stimulation (LU, right) conditions for sighted controls (SC, green, top) and individuals with reversed congenital cataracts (CC, red, bottom). Displayed p values are Bonferroni corrected for multiple comparisons.

Revised version: Correlations between water-normalized Glx concentration and aperiodic intercept are shown for the eyes closed (EC, left), eyes open (EO, middle) and visual stimulation (LU, right) conditions for sighted controls (SC, green, top) and individuals with reversed congenital cataracts (CC, red, bottom). The reported adjusted p values (p-adj) are Bonferroni corrected for multiple comparisons. The 95% confidence intervals (CI) of the correlation coefficients (r) are reported.

A negative correlation between the aperiodic slope and Glx concentration in CC individuals (i.e., steeper slopes with increasing Glx concentration) was observed during visual stimulation, but did not survive correction for multiple comparisons (Supplementary Material S17). No such correlation was observed between Glx concentration and aperiodic slope in the eyes open or closed conditions. Visual cortex GABA+ concentration and Glx/GABA+ concentration ratios did not significantly correlate with the aperiodic intercept or slope in either CC or SC individuals, during any experimental condition (Supplementary Material S17).

Discussion

Research in non-human animals has provided convincing evidence that the ratio of excitation to inhibition (E/I) in the visual cortex is reliant on early visual experience (Froemke, 2015; Haider et al., 2006; Hensch et al., 1998; Takesian & Hensch, 2013; Wu et al., 2022). Studies in humans who were born blind due to dense bilateral cataracts, and who had received sight restoration surgery in childhood or as adults, have found limited recovery of both basic visual and higher order visual functions (Birch et al., 2009; Röder & Kekunnaya, 2021). Here, we tested whether neurotransmitter concentrations and electrophysiological markers of cortical E/I ratio depend on early visual experience in humans, and how possible changes in visual cortex E/I ratio relate to sight recovery. First, we employed Magnetic Resonance Spectroscopy (MRS) and assessed Glutamate/Glutamine (Glx) and Gamma-Aminobutyric Acid (GABA+) concentrations, as well as their ratio, in the visual cortex (Shibata et al., 2017; Steel et al., 2020; Takei et al., 2016). Second, the slope and intercept of the aperiodic resting-state EEG activity with eyes open and eyes closed (R. Gao et al., 2017; Muthukumaraswamy & Liley, 2018; Ossandón et al., 2023), as well as during flickering visual stimulation (Pant et al., 2023), were measured over the occipital cortex in the same individuals. The EEG measures allowed us to exploratorily relate neurotransmitter changes to neural activity changes in congenital cataract-reversal individuals.

We found a lower Glx/GABA+ concentration ratio in the visual cortex of congenital cataract-reversal (CC) individuals as compared to normally sighted controls (SC). Additionally, the slope of the aperiodic EEG power spectrum was steeper for the low frequency range (1-20 Hz) and its intercept was higher in CC than SC individuals. In the CC group, Glx concentration correlated with the intercept of the aperiodic component during flickering visual stimulation. The Glx/GABA+ concentration ratio during the eyes closed condition predicted visual acuity of CC individuals. Together, the present results provide initial evidence for experience-dependent development of the E/I ratio in the human visual cortex, with consequences for behavior.

Previous MRS studies in the visual cortex of permanently congenitally blind humans reported higher Glx concentrations (Coullon et al., 2015) in five anophthalmic humans, and numerically lower GABA concentrations in congenitally blind humans (Weaver et al., 2013) (n = 9), as compared to normally sighted individuals. These results were interpreted as suggesting a higher E/I ratio in the visual cortex of permanently congenitally blind humans, which would be consistent with the extant literature on higher BOLD activity in the visual cortices of the same population (Bedny, 2017; Röder & Kekunnaya, 2022). We observed a lower Glx/GABA+ ratio and a steeper slope of the aperiodic EEG activity (1-20 Hz) at occipital electrodes, both of which suggest a lower rather than higher E/I ratio in the visual cortex of CC individuals. Here, we speculate that our results imply a change in neurotransmitter concentrations as a consequence of restoring vision following congenital blindness. Further, we hypothesize that due to limited structural plasticity after a phase of congenital blindness, the neural circuits of CC individuals, which had adapted to blindness after birth, likely employ physiological plasticity mechanisms (Knudsen, 1998; Mower et al., 1985; Röder et al., 2021), in order to re-adapt to the newly available visual excitation following sight restoration later in life.

Structural remodeling (Bourgeois, 1996) for typical E/I balance requires visual experience following birth (Hensch & Fagiolini, 2005; Takesian & Hensch, 2013; H. Zhang et al., 2018) and is linked to a sensitive period (Desai et al., 2002; Hensch & Fagiolini, 2005). A repeatedly documented finding in permanently congenitally blind humans is the increased thickness of visual cortex (Anurova et al., 2015; Hölig et al., 2023; J. Jiang et al., 2009). These structural changes in permanently congenitally blind individuals were interpreted as a lack of experience-dependent pruning of exuberant synapses and/or reduced myelination, the latter typically leading to a shift of the grey-white matter boundary (Natu et al., 2019). In parallel, it was observed in non-human primates that the overproduction of synapses during the initial phase of brain development was independent of experience, but that synaptic pruning, predominantly of excitatory synapses, depended on visual experience (Bourgeois, 1996; Bourgeois et al., 1989). The lack of excitatory synapse pruning was thought to underlie the observed higher excitability of visual cortex due to congenital visual deprivation (Benevento et al., 1992; Huang et al., 2015; Morales et al., 2002). Crucially, increased visual cortex thickness (Feng et al., 2021; Guerreiro et al., 2015; Hölig et al., 2023) and higher BOLD activity during rest with the eyes open (Raczy et al., 2022) have been observed for CC individuals following congenital blindness, suggesting incomplete recovery of cortical structure and function after sight restoration in humans. Thus, the restored feedforward drive to visual cortex after cataract removal surgery might reach a visual cortex with a lower threshold for excitation.

As to the best of our knowledge, the present study was the first assessing MRS markers of cortical excitation in humans following sight restoration after congenital blindness, we allude to non-human animal work for the interpretation of our findings. Such studies have often demonstrated that excitation and inhibition go hand-in-hand (Froemke, 2015; Haider et al., 2006; Isaacson & Scanziani, 2011; Tao & Poo, 2005). Analogously, we speculate that an overall reduction in Glx/GABA ratio might be effective in counteracting the aforementioned adaptations to congenital blindness, i.e. a lower threshold for excitation. Higher overall visual cortex excitation as a consequence of congenital blindness (Benevento et al., 1992; Morales et al., 2002) might come with the risk of runaway excitation in the presence of restored visually-elicited excitation. Phrased differently, we postulate that significantly lowering overall excitation of an originally hyper-excited visual cortex (during the phase of blindness) after sight restoration might be instrumental for maintaining neural circuit stability. Evidence for such homeostatic adjustments comes from studies with normally sighted humans that observed a reduction of GABA concentrations in visual cortex (Lunghi et al., 2015) and an increase in the BOLD response (Binda et al., 2018) following monocular blindfolding. Further, studies in adult mice have provided support for a homeostatic adjustment of the E/I ratio following prolonged changes in neural activity (Chen et al., 2022; Goel & Lee, 2007; Keck et al., 2017; Whitt et al., 2013). For example, a long period of decreased activity following enucleation in adult mice commensurately decreased inhibitory drive (Keck et al., 2011), primarily onto excitatory neurons (Barnes et al., 2015). In line with the lowered Glx/GABA+ ratio being a compensatory measure to prevent runaway excitation during visual stimulation, the link between the Glx/GABA+ ratio during eye closure and visual acuity in an exploratory correlation analysis suggests that the more successful such assumed downregulation of the E/I ratio in visual cortex, the better the visual recovery. In fact, this correlation with visual acuity recovery is reminiscent of a previously reported correlation in a larger group of CC individuals, between decreased visual cortex thickness and better visual acuity (Hölig et al., 2023). Hence, CC individuals with more advanced structural normalization appear to have a better starting point for functional recovery, the latter possibly mediated by physiological plasticity. Yet, future work has to explicitly test these hypotheses.

An increased intercept of the aperiodic component of occipital EEG activity was observed in the same CC individuals who underwent MRS assessment, irrespective of condition, that is, during rest with eyes open and eyes closed, as well as during flickering stimulation. The intercept of the aperiodic component has been linked to overall neuronal spiking activity (Manning et al., 2009; Musall et al., 2014) and fMRI BOLD activity (Winawer et al., 2013). The higher aperiodic intercept may therefore signal increased spontaneous spiking activity in the visual cortex of CC individuals. This interpretation would be consistent with the previously observed increase in visual cortex BOLD activity of CC compared to SC individuals (Raczy et al., 2022).

In CC individuals, the intercept of the aperiodic activity was highly correlated with the Glx concentration during rest with eyes open and during flickering stimulation. This exploratory finding needs replication in a larger sample. If reliable, the correlation between the EEG aperiodic intercept and Glx concentration in CC individuals might indicate more broadband firing (Manning et al., 2009; Winawer et al., 2013) in CC than SC individuals during active and passive visual stimulation.

Limitations

The sample size of the present study was rather high for rare population of carefully diagnosed CC individuals, but undoubtedly overall small. Access to CC individuals was limited by the constraints of the COVID-19 pandemic. Hence, all the group differences, the exploratory correlations with visual history metrics, and between MRS-EEG parameters, are reported for further investigation in a larger sample. Moreover, our speculative accounts for the present findings need to be validated with pre- and post-surgery assessments. Finally, a comparison of CC individuals with a control group of developmental cataract-reversal individuals would be instrumental to test the hypothesis that the observed group differences are specific to early brain development.

We are aware that MRS and EEG has a low spatial specificity. Moreover, MRS measures do not allow us to distinguish between presynaptic, postsynaptic and vesicular neurotransmitter concentrations. However, all reported group differences in MRS and EEG parameters were specific to visual cortex and were not found for the frontal control voxel or at frontal electrodes, respectively. While data quality was lower for the frontal compared to the visual cortex voxels, as has been observed previously (Juchem & Graaf, 2017; Rideaux et al., 2022), this was not an issue for the EEG recordings. Thus, lower sensitivity of frontal measures cannot easily explain the lack of group differences for frontal measures. Crucially, data quality did not differ between groups.

While interpretations of new data in the absence of similar data sets are necessarily speculative, the validity of the neurochemical findings was supported by quality assessments; phantom testing showed high correlations between the experimentally varied metabolite concentrations and the extracted GABA+ and Glx concentrations (Supplementary Figure S3). The neurochemical results were robust to analysis pipelines (Supplementary Material S3) as well as normalization method (Supplementary Material S5). The EEG results from the present group of CC individuals replicated effects observed in a larger sample of 28 additional CC individuals (Supplementary Material S18) (Ossandón et al., 2023), as well as prior findings from another sample reporting lower alpha power (Supplementary Material S19) (Bottari et al., 2016; Pant et al., 2023). Further, the aperiodic intercept of EEG activity decreased with chronological age irrespective of group or condition, replicating earlier reports (Hill et al., 2022; Voytek et al., 2015) (Supplementary Material S15). Finally, group differences were observed despite the considerable variance of blindness duration and time since surgery, demonstrating the crucial role of early visual experience.

Conclusion

The present study in sight recovery individuals with a history of congenital blindness indicates that E/I balance is a result of early experience and crucial for human behavior. We provide initial evidence that the E/I ratio in congenital cataract-reversal individuals is altered even years after surgery, which may be due to previous adaptation to congenital blindness.